// codefestival_2016_qualB_b: B - Qualification simulator
// https://atcoder.jp/contests/code-festival-2016-qualb/tasks/codefestival_2016_qualB_b

#include <bits/stdc++.h>

using namespace std;

int main()
{
    int n, a, b;
    string s;
    cin >> n >> a >> b >> s;
    // I AM NOT DONE
    return 0;
}